# DATA 601

## HW1


**Learning Objectives**

- Explore built-in data types in Python.
- Review fundamental programming and problem solving concepts with Python.
- Implement functions based on mathematical concepts and definitions.
- Gain experience working with the Jupyter notebook environment.

_This is an individual homework assignment._ 

Please complete this homework assignment within the Jupypter notebook environment.  

#### Submission 

In order to ensure that the submission and grading process goes smoothly, please follow these instructions:

- Please provide your solutions where asked; please do not alter any other parts of this notebook.

- Save your notebook in its original IPYNB format and submit it to the dropbox folder.


```python
# Check that we are using a recent version of Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Part A

This part focuses on scalar types. You should be able to complete the following questions without using any collection types.

### Question 1

The central bionomial coefficients form a vertical row down the middle of Pascal's triangle. They can be calculated via the following formula:

$$
{2n\choose n} = \frac{(2n)!}{(n!)^2} = \prod_{k=1}^{n} \frac{n+k}{k}  
$$

Write a function that calculates the nth central bionomial coefficient. You may use in-built functions to calculate the factorials.



```python
def cent_binomial_coef(n):
    '''Returns the n-th binomial coefficient for an input
    integer n.
    '''
    ...
```

##### Question 2



Many things in nature can be modeled by the exponential growth function: 
$$
P(t) = \alpha \beta^{\frac{t}{\tau}},
$$
where $\alpha$ is the initial population, $\beta$ is the growth rate factor, t is the time passed, and $\tau$ is the time constant. 

Imagine that you are an entomologist on the plant Zergo. There is a small insect of interest to you on this planet which has an exponential growth rate. From experimentation you determine that: 
$\beta= \sqrt\frac{5}{4}$,  and $\tau= 2^{\frac{1}{3}}$.

Starting with a population $\alpha$ of 10, use this formula to compute $P(t)$ in the function `population(t)` below. 


```python
import math
def population(t):
    '''Computes the population at time t, where t is a positive float, 
    and returns the population as an integer (floor value). 
    It is assumed that time is in minutes and that all parameters
    are correct for this assumption.
    '''
    ...
```

## Part B

Questions in this part rely on strings, lists, and tuples.

### Question 1

Consider the lists $a=[2,1,2,5,3,1]$ and $b=[2,1,7,1]$. The intersection of these lists, $a \cap b$, is $[1,1,2]$. Write a function that returns the intersection of two arbitrary lists. You may not use built-in set operations, and should pay attention to duplicates.

(5 points)


```python
def intersection_list(a,b):
    '''
    Returns the intersection of two lists a and b. 
    
    '''
    ...
```

### Question 2


The scalar triple product of three vectors in $\mathbb{R^3}$ is defined as:

$$
a.(b \times c)
$$

In the above definition, '$.$' refers to a dot (scalar) product of two vectors in $\mathbb{R}^3$ while '$\times$' denotes a cross product of two vectors in $\mathbb{R}^3$.

Write a function that calculates the scalar triple product of three vectors, $a$, $b$, and $c$ that are in $\mathbb{R}^3$.

```python
def striple_prod(a,b,c):
    '''Returns the scalar triple product of three vectors A, B, and C. 
    '''
    ...
```

### Question 3a

Write a function that calculates the most frequently occurring letter given a piece of text supplied as a string.


```python
def freq_letters(text):
    '''
    For textual input, returns the most frequently occuring letter 
    (your function should be case-insensitive)
    '''
    ...
```

### Question 3b


Write a second function that returns the number of vowels and consonants in the text.

```python
def vowels(text):
    '''
    For textual input, returns the number of vowels and consonants.
    Return a tuple with number of vowels as the first element and consonants
    as the second element, i.e. (num_vowels, num_consonant)
    '''
    ...
```

## Part C

### Tasks

#### 1. Data Loading
- Read and load the dataset (spam.txt) into your program using Python's built-in libraries or by manually processing the text file.
- Split the data into two arrays:
  - **Labels**: Indicating whether a message is 'spam' or 'ham'.
  - **Messages**: Containing the actual text of the messages.

#### 2. Data Exploration
- Calculate:
  - The total number of messages in the dataset.
  - The number of messages labeled as 'spam' and 'ham'.
  - The percentage of spam messages in the dataset.

#### 3. Text Analysis
- Perform the following:
  - Find the average number of words per message for both 'ham' and 'spam'.
  - Identify the longest message (in terms of word count) in each category ('spam' and 'ham') and display it.

#### 4. Data Visualization
- Use Matplotlib to create the following visualizations:
  - **Bar Chart**: Display the number of messages in each category ('spam' vs. 'ham').
  - **Histogram**: Show the distribution of message lengths (word counts) for 'spam' and 'ham' messages.
  - **Pie Chart**: Represent the proportion of 'spam' and 'ham' messages in the dataset.

#### 5. Further Analysis
- Perform further analysis using Numpy:
  - Identify the top 10 most frequently occurring words in the 'spam' messages along with their counts.
  - Plot a bar chart showing the frequency of these top 10 words.

#### 6. Reflection
- Write a short paragraph reflecting on:
  - The challenges faced while analyzing the data.
  - Insights gained about spam messages from the analysis.


_Type your answer here, replacing this text._

```python

```

Tests for q1_a (points: 3)
```
>>> cent_binomial_coef(0) == 1
True
```
failure_message: Don't forget (0)! is 1.
```
>>> cent_binomial_coef(1) == 2
True
```


Tests for q1_b (points: 5)
```
>>> def testcase1(intersection_list):
...     temp = intersection_list([2, 1, 2, 5, 3, 1], [2, 1, 7, 1])
...     length = len(temp)
...     ones = temp.count(1)
...     twos = temp.count(2)
...     assert (length, ones, twos) == (3, 2, 1)
>>> testcase1(intersection_list)
```


Tests for q2_a (points: 5)
```
>>> is_prime(2) == True
True
```
```
>>> is_prime(3) == True
True
```
```
>>> is_prime(4) == False
True
```


Tests for q2_b (points: 3)
```
>>> striple_prod([5, 1, 2], [-1] * 3, [2, 1, 0]) == 5
True
```


Tests for q3_a (points: 3)
```
>>> population(5.7) == 16
True
```
```
>>> population(2.9) == 12
True
```


Tests for q3_b_a (points: 5)
```
>>> freq_letters('a!!!!!').lower() == 'a'
True
```
failure_message: You should only count letters a-z and not numbers, punctuation, spaces, etc.
success_message: Good job!


Tests for q3_b_b (points: 4)
```
>>> vowels('I like banana!!') == (6, 5)
True
```
